# Assignment 1

For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"

    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    # raise NotImplementedError()
```

## Part A

Find a list of all of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re


def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""
    pattern = "[A-Z][a-z]*"
    result = re.findall(pattern, simple_string)
    return result
    names()
    raise NotImplementedError()
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"
```

```python

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re


def grades():
    with open("assets/grades.txt", "r") as file:
        grades = file.read()

    # Regex pattern to find names of students who received a B
    pattern = r"([A-Za-z ]+): B\b"

    # Find all matches in the string
    matches = re.findall(pattern, grades)

    return matches


# Testing the function
print(grades())  # This will print the list of students who received a B
assert len(grades()) == 16  # Ensure there are 16 students with a B grade
```

Output:
```
[]
AssertionError                            Traceback (most recent call last)
Cell In [4], line 13
      8     return matches
     10 print(grades())
---> 13 assert len(grades())==16

AssertionError:
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16
```

Output:
```
Cell In [24], line 1
----> 1 assert len(grades()) == 16

Cell In [23], line 12, in grades()
      9             grades()
     10             len(grades())
---> 12 raise NotImplementedError()

NotImplementedError:
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re


def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()

    # Regex pattern to capture host, username, time, and request
    pattern = r'(?P<host>\S+) - (?P<user_name>\S+) \[(?P<time>.*?)\] "(?P<request>.*?)"'

    # Find all matches in the log data
    matches = re.finditer(pattern, logdata)

    # Convert matches to list of dictionaries
    log_list = []
    for match in matches:
        log_list.append(match.groupdict())

    return log_list


# Testing the function
print(logs())  # This will print the list of dictionaries
assert len(logs()) == 979  # Ensure there are 979 log entries

one_item = {
    "host": "146.204.224.152",
    "user_name": "feest6811",
    "time": "21/Jun/2019:15:45:24 -0700",
    "request": "POST /incentivize HTTP/1.1",
}

assert (
    one_item in logs()
), "Sorry, this item should be in the log results, check your formatting"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In [30], line 27
     19 # Test assertions
     20 one_item = {
     21     'host': '146.204.224.152',
     22     'user_name': 'feest6811',
     23     'time': '21/Jun/2019:15:45:24 -0700',
     24     'request': 'POST /incentivize HTTP/1.1'
     25 }
---> 27 assert len(logs()) == 979, "The number of log entries does not match the expected value."
     28 assert one_item in logs(), "Sorry, this item should be in the log results, check your formatting."

AssertionError: The number of log entries does not match the expected value.
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item = {
    "host": "146.204.224.152",
    "user_name": "feest6811",
    "time": "21/Jun/2019:15:45:24 -0700",
    "request": "POST /incentivize HTTP/1.1",
}
assert (
    one_item in logs()
), "Sorry, this item should be in the log results, check your formating"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In [31], line 1
----> 1 assert len(logs()) == 979
      3 one_item={'host': '146.204.224.152',
      4   'user_name': 'feest6811',
      5   'time': '21/Jun/2019:15:45:24 -0700',
      6   'request': 'POST /incentivize HTTP/1.1'}
      7 assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

AssertionError:
```

```python

```